---
title       : Testing SQL
description : Testing SQL

--- type:NormalExercise xp:100 key:c778ff1b1a
## SQL NormalExercise

This is a SQL normal exercise.

*** =instructions
- Just experiment!

*** =pre_exercise_code
```{python}
# this is python code
```

*** =sample_code
```{sql}
# this is sql code
```

*** =solution
```{sql}
# this is sql code
```

*** =sct
```{python}
# this is python code
success_msg('awesome!')
```
